*Hello, this is a test*

Time to find out what students can see. 

*Exercise*

Assignment $A$ the value 42. 

<!-- grade_id: cell-00fd7b523f55298e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a7c4c61000927e59, points: 0, grade: true, solution: false -->
```python
assert isinstance(A,int), "NOT AN INT"
print("Great int")
```

<!-- grade_id: cell-5df41029913dcc6c, points: 0, grade: true, solution: false -->
```python

```

```python

```