# Assignment 0: Getting to know the workflow

Welcome to Math 307, specifically the python workbook component. This aparatus is new to us as well as to you! We expect a few growing-pains at the start of term as we all familiarize ourselves with the system. We also expect that once we're all in the swing of things, these notebooks will afford us a way to give you A LOT of feedback really QUICK! 

A reminder that every Wednesday at 6pm PST, Ethan White will run a python office hour to answer all questions related to the workflow of these notebooks. This first notebook is not for marks, but we highly recommend everyone complete it so that we find identify any problems when there are no marks on the line :) 

```python
import numpy as np
```

**Exercise 1**

Create a variable $C$ and assign it the value 307.

<!-- grade_id: cell-6a9b4111a311fc60, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

After we ask you all a question, we'll perform multiple 'tests' on your answer to award as many part marks as possible, and provide as much feedback as possible. 

Feel free to deliberately answer this question wrong to trigger the different feedbacks. 

<!-- grade_id: cell-ca45778045f0107c, points: 1, grade: true, solution: false -->
```python
assert isinstance(C,int), "Your C is not an integer"
print("Checkpoint 1.1 passed!")
```

<!-- grade_id: cell-ff6a2ba9f355f1af, points: 1, grade: true, solution: false -->
```python
assert np.isclose(C,307), "You didn't give the right value to C"
print("Checkpoint 1.2 passed!")
```

**Exercise 2**

Create the matrix 

$$\begin{pmatrix} 3 & -5 \\ 2 & 1.6 \end{pmatrix} $$

and assign it to the variable $M$. 

<!-- grade_id: cell-8bb67e8f6990f9f1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

Once again, feel free to make $M$ be the wrong answer to trigger some of the feedback as you complete the assignment... and then change it to the correct answer :) 

For this question, we will check that $M$ is an array for one mark, and that the array is $2 \times 2$ for one mark, and we will also do a final 'hidden test' for two marks that checks that the whole matrix is correct. 

In assignments appearing throughout this term, some of the questions we give will have feedback you can immediately see (as you are completing the assignment), and some will have hidden tests, where you wont know if you got the right answer until after you have submitted the assignment back in Canvas, and we have graded it. 

<!-- grade_id: cell-2ce390a29bbb3425, points: 1, grade: true, solution: false -->
```python
assert isinstance(M,np.ndarray), "M needs to be an array. Remember to import numpy"
print("Checkpoint 2.1 passed!")
```

<!-- grade_id: cell-691b5a8c4460c786, points: 1, grade: true, solution: false -->
```python
assert M.shape == (2,2), "The dimensions of your array are incorrect."
print("Checkpoint 2.2 passed!")
```

<!-- grade_id: cell-999dc6e4cf7ae8c5, points: 2, grade: true, solution: false -->
```python

```

*Exercise 3*

Rank the following four numbers from smallest to largest. 

a) $\pi$
b) $2\sqrt{3}$
c) $\frac{22}{7}$
d) $3.14$

Record your answer as a python list $L$, where each entry is a single-character string, for example ['a','b','c','d']. Remember to use quotes '' to create a string, rather than simply entrying the letter c. This example list is not in the correct order!

<!-- grade_id: cell-dc4814a9e89c49e0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-6853910a118a2d3d, points: 1, grade: true, solution: false -->
```python
assert isinstance(L,list), "Make sure L is a list, created with square brackets []."
print("Checkpoint 3.1 passed!")
```

<!-- grade_id: cell-1fe91d961d08c0a7, points: 1, grade: true, solution: false -->
```python
assert len(L) == 4, "Make sure you list has length 4, one entry for 'a','b','c','d'"
print("Checkpoint 3.2 passed!")
```

<!-- grade_id: cell-cdf2984714a88b93, points: 1, grade: true, solution: false -->
```python
assert L[0] == 'd' and L[3] == 'b', "You haven't identified the largest and smallest numbers correctly."
print("Checkpoint 3.3 passed!")

```

<!-- grade_id: cell-30f8c0adb31462ab, points: 2, grade: true, solution: false -->
```python

```

Once again, there was a final test on $L$ that was hidden from you as you complete the assignment. Once your assignment has been returned to you, you'll find out if you ranked all the numbers in the correct order. 